# EEN060 - Applied object-oriented programming

Teacher: [Carlos Natalino](https://www.chalmers.se/en/staff/Pages/Carlos-Natalino-Da-Silva.aspx) / Examiner: [Paolo Monti](https://www.chalmers.se/en/staff/Pages/Paolo-Monti.aspx)

[Canvas course page](https://chalmers.instructure.com/courses/28094)

[Course channel on Chalmers Play](https://play.chalmers.se/channel/EEN060_EEN065_Applied_object-oriented_programming/300149)

Before you turn this assignment list in, make sure everything runs as expected.
First, **restart the kernel** and then **run all cells**.
Then, check if all the tests run correctly.
Note that if one of the problems present an error, the following ones **will not** be tested.

In case of discrepancies between the problem command and the tests, you should solve it having in mind the tests.

There are two types of cell:
1. *solution cells:* These are the cells where you write your answer, or modify the existing code to solve the problem.
2. *test cells:* These cells are used to test whether your solution is correct or not. If the tests run correctly, you should see a message `tests passed`. Otherwise, you should see an error message.

**Delete** the line `raise NotImplementedError()` from the problems that you solve.

**Do not delete or add any cell in this file.** All cells that you need are already in place.

You may also have problems that will be manually graded (e.g., HTML/CSS/JS/pseudocode problems).
These problems **do not** have a respective *test cell*.

If you want to execute a cell, select the cell and press **CTRL+Enter** (in Windows) or **CMD+Enter** (in macOS) or click on the **Run cell** button.

**If you want, you can solve this programming assignment using Google Colab**

Link: https://colab.research.google.com/

Just copy the following line to a code line and run it.

```python
# !pip install --upgrade --user pip pip-tools
# !pip install ipykernel jupyter notebook wtforms flask flask-wtf gunicorn nose flask-testing pytest coverage flake8 requests mypy pylint blinker types-requests isort black lorem-text bs4 pep8-naming matplotlib redis
```

# Assignment 4

In this assignment, students should practice:
 - how to use built-in functionalities
 - how to use standard Python modules
 - how to find and use external Python modules
 - how to define classes
 - how to instantiate and intialize objects
 - how to define attributes
 - how to define methods
 - how to use inheritance
 - how to use composition
 - how to use dataclasses
 - how to set HTML styles property

 **Observations:**
 For the problems in this assignment, you should:
 - use the command `import`, if needed, within the solution cell, every time you need to use some module.
- make sure that your `utils.py` file has the following line (i.e., replace `onu2` by `onu1` if needed):

`url_html_validator = "https://onu1.s2.chalmers.se/nu/?out=json"`

 Reading material:
 - lecture notes
 - Online documentation:
   - Python online documentation: https://docs.python.org/3.10/tutorial/modules.html
   - Specific modules (for optional reading):
     - datetime: https://docs.python.org/3.10/library/datetime.html
     - math: https://docs.python.org/3.10/library/math.html
     - random: https://docs.python.org/3.10/library/random.html
     - statistics: https://docs.python.org/3.10/library/statistics.html

**Preparation:** Run the cell below every time you start working on this file, and every time you restart the kernel.

```python
%load_ext autoreload
try:
    from utils import validate_python_code, validate_html
except:
    import sys
    print("It seems this file is in the wrong folder. "
          "Make sure to place it in the `programming-assignments` folder/project.",
          file=sys.stderr)
```

## Level 1

Level 1 problems are simple problems that required only the knowledge obtained in the content of the lecture and the lecture notes of the particular lecture.
In our case, the official documentation of Python also applies.

### Problem 1.1

Create a Python function called `prod_list` that receives a list of numbers (`int` or `float`) as an input, and returns the product of all the elements in the list.

You must use the function `prod` from `math` module to solve this problem.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-1, points: 0, grade: false, solution: true -->
```python
%%writefile prod_list_solution.py
# solution cell
### BEGIN SOLUTION
from math import prod


def prod_list(lst: list[int | float]) -> int | float:
    return prod(lst)


### END SOLUTION

```

Output:
```
Overwriting prod_list_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from prod_list_solution import prod_list

input_list = [10, 12.0, 15, 17] 
product = prod_list(input_list)
print(f'The product of {input_list} is {product}')
print('Execution finished', u'\u2713')
```

Output:
```
The product of [10, 12.0, 15, 17] is 30600.0
Execution finished ✓
```

<!-- grade_id: cell-tests-1-1, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import prod_list_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from prod_list_solution import prod_list
except:
    raise ValueError("Your solution does not contain the right function!")

with open("prod_list_solution.py", "rt", encoding="UTF-8") as file:
    content = file.read()
    assert "math" in content, "You must import the `math` module."

from nose.tools import assert_equal

assert_equal(prod_list([]), 1)
assert_equal(prod_list([1, 2, 3.0, 4, 5]), 120.0)
assert_equal(prod_list([10, 12.0, 15, 17]), 30600.0)
assert_equal(prod_list([10, 12.0, 15, -17]), -30600.0)

validate_python_code("prod_list_solution.py")
print('tests passed', u'\u2713')
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 1.2

Create a Python function called `gm_median` that receives a list of numbers (`int` or `float`) as parameter.
The function should return the geometric mean, and the median of the list, computed using the module `statistics`. The result should be returned as a `tuple`.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-2, points: 0, grade: false, solution: true -->
```python
%%writefile gm_median_solution.py
# solution cell
### BEGIN SOLUTION
import statistics


def gm_median(lst: list[int | float]) -> tuple[int | float, int | float]:
    # Calculate geometric mean
    gm = statistics.geometric_mean(lst)
    # Calculate the median
    median_value = statistics.median(lst)
    return gm, median_value


### END SOLUTION

```

Output:
```
Overwriting gm_median_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from gm_median_solution import gm_median

numbers = [30.36, 90.42, 35, 42.3, 23.0, 21.1, 90, 39.63, 60.46, 62.1]
print(f'Numbers: {numbers}')
stats = gm_median(numbers)
print(f'Geomtric Mean: {stats[0]}')
print(f'Median: {stats[1]}')
# print('Execution finished', u'\u2713')
```

Output:
```
Numbers: [30.36, 90.42, 35, 42.3, 23.0, 21.1, 90, 39.63, 60.46, 62.1]
Geomtric Mean: 43.87136324804066
Median: 40.965
```

<!-- grade_id: cell-test-1-2, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import gm_median_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from gm_median_solution import gm_median
except:
    raise ValueError("Your solution does not contain the right function!")

with open("gm_median_solution.py", "rt", encoding="UTF-8") as file:
    content = file.read()
    assert "statistics" in content, "You must import the `statistics` module."

from nose.tools import assert_equal

assert_equal(isinstance(gm_median([15.88, 48.0, 53.9, 4, 93]), tuple), True)
assert len(gm_median([15.88, 48.0, 53.9, 4, 93])) == 2, 'The return of the function must be a sequence with 3 elements.'
assert_equal(gm_median([15.88, 48.0, 53.9, 4, 93]), (27.34290717539487, 48.0))
assert_equal(gm_median([2, 69, 39, 40, 35, 99]), (30.115112319375513, 39.5))

from unittest.mock import patch

with patch('statistics.geometric_mean') as mock_stdev:
    gm_median([2.58, 69.7])
mock_stdev.assert_called_once()

validate_python_code("gm_median_solution.py")
print('tests passed', u'\u2713')
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 1.3

Create a Python function called `random_mobile_number_generator` that randomly generates a mobile number (`str`). 

The number should start with either of `0`, `46`, or `+46`. The remaining part of the number has length 9 such that the first digit is between 1 to 9 and the other digits are between 0 to 9.

You must import and use module `random` in your solution.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-3, points: 0, grade: false, solution: true -->
```python
%%writefile random_mobile_number_solution.py
# solution cell
### BEGIN SOLUTION
import random


def random_mobile_number_generator() -> str:
    # Randomly choose the initial part of the number
    prefix_options = ["0", "46", "+46"]
    prefix = random.choice(prefix_options)

    # Generate the remaining 9 digits
    first_digit = random.randint(1, 9)  # The first digit must be between 1 and 9
    remaining_digits = [str(random.randint(0, 9)) for _ in range(8)]
    number_body = str(first_digit) + "".join(remaining_digits)
    # Combine the prefix with the number body
    mobile_number = prefix + number_body
    return mobile_number


### END SOLUTION

```

Output:
```
Overwriting random_mobile_number_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from random_mobile_number_solution import random_mobile_number_generator

print('Here are some random mobile numbers:')
for i in range(5):
    print(random_mobile_number_generator())
print('Execution finished', u'\u2713')
```

Output:
```
Here are some random mobile numbers:
+46322311650
+46981166586
+46382939286
+46963886865
+46857737755
Execution finished ✓
```

<!-- grade_id: cell-test-1-3, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import random_mobile_number_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from random_mobile_number_solution import random_mobile_number_generator
except:
    raise ValueError("Your solution does not contain the right function!")

with open("random_mobile_number_solution.py", "rt", encoding="UTF-8") as file:
    content = file.read()
    assert "random" in content, "You must import the `random` module."

from nose.tools import assert_almost_equal
# testing the probability
import collections
import random

idx_digit = random.randint(1, 8)
number_samples = 100000
count_prefix = collections.defaultdict(int)
count_digit = collections.defaultdict(int)
valid_prefixes = ["0", "46", "+46"]

for i in range(number_samples):
    number = random_mobile_number_generator()
    res = [number.startswith(prefix) for prefix in valid_prefixes]
    assert any(res) == True, f"Generated number {number} is invalid due to incorrect prefix. Check the rules again."
    prefix = valid_prefixes[res.index(True)]
    count_prefix[prefix] += 1
    remaining_part = number.replace(prefix, '', 1)
    assert remaining_part.isdigit() == True, f"Generated number {number} is invalid due to incorrect digits. Check the rules again."
    assert len(remaining_part) == 9, f"Generated number {number} is invalid due to incorrect length. Check the rules again."
    assert remaining_part[0] != '0', f"Generated number {number} is invalid due to 0 in first digit after prefix. Check the rules again."
    count_digit[remaining_part[idx_digit]] += 1

for _, count in count_prefix.items():
    assert_almost_equal(count / number_samples, 1/3, places=2)

for _, count in count_digit.items():
    assert_almost_equal(count / number_samples, 1/10, places=2)

validate_python_code("random_mobile_number_solution.py")
print('tests passed', u'\u2713')
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 1.4

Create a Python class with the name `VolumeCalculator`.

The initialization `.__init__()` method receives a number and assign the value to the attribute `radius`.

Then, create a method named `volume()` that returns the volume of the sphere using the `radius` attribute.
If the `radius` attribute is `None` or $\leq 0$ the method `volume()` should return `None`.

**Remember:** from the `math` module, import `pi`.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: answer-1-5, points: 0, grade: false, solution: true -->
```python
%%writefile volume_calculator_class.py
# solution cell
### BEGIN SOLUTION
from math import pi


class VolumeCalculator:
    def __init__(self, radius: float | int | None) -> None:
        self.radius = radius

    def volume(self) -> float | None:
        if self.radius is None or self.radius <= 0:
            return None
        return (4 / 3) * pi * self.radius**3


### END SOLUTION

```

Output:
```
Overwriting volume_calculator_class.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from volume_calculator_class import VolumeCalculator

radius = 2
s = VolumeCalculator(radius)
print(f'The volume of a sphere with radius {radius} is {s.volume()}')
print('Execution finished', u'\u2713')
```

Output:
```
The volume of a sphere with radius 2 is 33.510321638291124
Execution finished ✓
```

<!-- grade_id: tests-1-5, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import volume_calculator_class
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from volume_calculator_class import VolumeCalculator
except:
    raise ValueError("Your solution does not contain the right function!")

with open("volume_calculator_class.py", "rt", encoding="UTF-8") as file:
    content = file.read()
    assert "math" in content, "You must import `pi` from the `math` module."

from nose.tools import assert_equal, assert_almost_equal
import math

assert 'VolumeCalculator' in dir(), 'Class VolumeCalculator does not exist!'
assert 'volume' in dir(VolumeCalculator), 'Method volume() does not exist!'
s = VolumeCalculator(None)
assert s.volume() == None, 'Method volume() should return None for invalid values.'
s = VolumeCalculator(0)
assert s.volume() == None, 'Method volume() should return None for invalid values.'
s = VolumeCalculator(-1)
assert s.volume() == None, 'Method volume() should return None for invalid values.'
s = VolumeCalculator(1)
assert isinstance(s.volume(), float) or isinstance(s.volume(), int), 'The return type of the function volume() does not match the problem description.'
assert_almost_equal(s.volume(), 4.18, places=1)
s = VolumeCalculator(2)
assert isinstance(s.volume(), float) or isinstance(s.volume(), int), 'The return type of the function volume() does not match the problem description.'
assert_almost_equal(s.volume(), 33.5, places=1)

validate_python_code("volume_calculator_class.py")
print('tests passed', u'\u2713')
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 1.5

Create a Python class named `Person`.
In the initialization method, the parameters are:
1. Name (mandatory)
2. role (mandatory)
3. person number (mandatory)
4. address (optional, `None` by default)
5. phone number (optional, `None` by default)

The initialization method should save these parameters as attributes.

The class should also have a `__repr__()` method that returns a string containing only the attributes that are not `None`, separated by `\t` characters.

```python
%%writefile person_class.py
# solution cell
### BEGIN SOLUTION
class Person:
    def __init__(
        self,
        name: str,
        role: str,
        person_num: str,
        address: str | None = None,
        phone_num: str | None = None,
    ) -> None:
        self.name = name
        self.role = role
        self.person_num = person_num
        self.address = address
        self.phone_num = phone_num

    def __repr__(self) -> str:
        # Collect attributes that are not None into a list
        attributes = [
            f"{self.name}",
            f"{self.role}",
            f"{self.person_num}",
        ]
        if self.address is not None:
            attributes.append(f"{self.address}")
        if self.phone_num is not None:
            attributes.append(f"{self.phone_num}")
        # Join the attributes with tab characters
        return "\t".join(attributes)


### END SOLUTION

```

Output:
```
Overwriting person_class.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from person_class import Person

person = Person('John Doe', 'Manager', '201414202014')
print(f'The Person object converts to string as {person}')
print('Execution finished', u'\u2713')
```

Output:
```
The Person object converts to string as John Doe	Manager	201414202014
Execution finished ✓
```

```python
%autoreload 2
# test cell
try:
    import person_class
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from person_class import Person
except:
    raise ValueError("Your solution does not contain the right function!")

from nose.tools import assert_equal
from unittest.mock import patch

assert 'Person' in dir(), 'Class Person does not exist!'
p = Person('John Doe', 'Manager', '201414202014' )

with patch('__main__.print') as mock_print:
    s = str(p)
mock_print.assert_not_called()

assert_equal(str(p), 'John Doe\tManager\t201414202014')
p = Person('Jane Doe', 'Supervisor', '201212202012', 'Gibraltargatan')
assert_equal(str(p), 'Jane Doe\tSupervisor\t201212202012\tGibraltargatan')
p = Person('Jane Doe', 'Supervisor', '201212202012', 'Gibraltargatan', '123456')
assert_equal(str(p), 'Jane Doe\tSupervisor\t201212202012\tGibraltargatan\t123456')
p = Person('Jane Doe', 'Supervisor', '201212202012', '123456')
assert_equal(str(p), 'Jane Doe\tSupervisor\t201212202012\t123456')
p = Person('Jane Doe', 'Super visor', '201212 202012',  'Gibraltar gatan')
assert_equal(str(p), 'Jane Doe\tSuper visor\t201212 202012\tGibraltar gatan')

validate_python_code("person_class.py")
print('tests passed', u'\u2713')
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

---
## Level 2

Problems in level 2 should build upon the content of the lecture and the content of previous lectures.
Therefore, the problems can be solved based solely on the lecture notes available for the course up to the lecture the assignment is about.

### Problem 2.1

Create a Python function called `lottery_numbers` that receives three parameters:
- how many numbers should be selected.
- the upper bound for the number generation;
- the lower bound for the number generation;

The function should generate a list of integer numbers with the correct number of elements, each element between the informed lower and upper bounds.
Remember that numbers in the list must not repeat.

The function should also validate the input. If any of the parameters is out of the range the function should return `None`. The conditions for parameters are:
- lower bound $\geq 0$, 
- upper bound $\geq 1$, 
- number $\ge 1$ 
- (upper bound - lower bound) $\geq$ (number - 1).


```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-1, points: 0, grade: false, solution: true -->
```python
%%writefile lottery_numbers_solution.py
# solution cell
### BEGIN SOLUTION
import random


def lottery_numbers(
    selected_num: int, upper_bound: int, lower_bound: int
) -> list[int] | None:

    # Validate inputs
    if lower_bound < 0:
        return None
    if upper_bound < 1:
        return None
    if selected_num < 1:
        return None
    if (upper_bound - lower_bound) < (selected_num - 1):
        return None

    # Generate unique random numbers within the range
    return random.sample(range(lower_bound, upper_bound + 1), selected_num)


### END SOLUTION

```

Output:
```
Overwriting lottery_numbers_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from lottery_numbers_solution import lottery_numbers

result = lottery_numbers(5, 50, 10)
print(f'Generating 5 numbers between 10 and 50 is {result}')
print('Execution finished', u'\u2713')
```

Output:
```
Generating 5 numbers between 10 and 50 is [30, 20, 48, 14, 47]
Execution finished ✓
```

<!-- grade_id: cell-test-2-1, points: 1, grade: true, solution: false -->
```python
# test cell
%autoreload 2
try:
    import lottery_numbers_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from lottery_numbers_solution import lottery_numbers
except:
    raise ValueError("Your solution does not contain the right function!")

from nose.tools import assert_equal

result = lottery_numbers(5, 50, 10)

assert len(result) == len(set(result)), 'There are repeated elements in the list'
assert min(result) >= 10
assert max(result) <= 50
assert len(result) == 5

result = lottery_numbers(30, 100, 0)

assert len(result) == len(set(result)), 'There are repeated elements in the list'
assert min(result) >= 0
assert max(result) <= 100
assert len(result) == 30

assert_equal(lottery_numbers(30, 100, -1), None)
assert_equal(lottery_numbers(30, 1, 0), None)

result = lottery_numbers(2, 1, 0)

assert len(result) == len(set(result)), 'There are repeated elements in the list'
assert min(result) >= 0
assert max(result) <= 1
assert len(result) == 2

assert_equal(lottery_numbers(0, 1, 0), None)
assert_equal(lottery_numbers(30, -1, 0), None)
assert_equal(lottery_numbers(-2, 10, 0), None)
assert_equal(lottery_numbers(0, 1, 0), None)

validate_python_code("lottery_numbers_solution.py")
print('tests passed', u'\u2713')
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 2.2

Create a Python dataclass named `PersonAge`.
The attributes of this class are:
1. name (Type: string)
2. job (Type: string)
3. birthdate (Type: date)
4. age (Type: int, not an argument during initialization)

The attributes `name`, `job` and `birthdate` are mandatory during initialization.
Then, you need to create a method `.compute_age()` that uses the `birthdate` to compute how old the person is, saves it to the `age` property, and returns the `age` value.

Remember to import `dataclass` and `field` from `dataclasses` module. Also import `datetime` for birthdate attribute.

Link to read more about data class: https://docs.python.org/3.10/library/dataclasses.html

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: solution-2-3, points: 0, grade: false, solution: true -->
```python
%%writefile person_age_class.py
# solution cell
### BEGIN SOLUTION
from dataclasses import dataclass, field
from datetime import date, datetime


@dataclass
class PersonAge:
    name: str
    job: str
    birthdate: date

    def compute_age(self) -> int:
        """Computes the age of the person based on the birthdate."""
        today = date.today()
        # Calculate age by subtracting birth year from current year
        calculated_age = today.year - self.birthdate.year
        # Adjust age if the birthdate hasn't occurred yet this year
        if (today.month, today.day) < (self.birthdate.month, self.birthdate.day):
            calculated_age -= 1
        # add the age attribute
        self.age = calculated_age
        return self.age


### END SOLUTION

```

Output:
```
Overwriting person_age_class.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from datetime import date
from person_age_class import PersonAge

person1 = PersonAge('John Doe', 'Manager', date(1990, 11, 15))
person2 = PersonAge('Jane Doe', 'Manager', date(1980, 11, 15))
print(f"The age of {person1.name} is {person1.compute_age()}")
print(f"The age of {person2.name} is {person2.compute_age()}")
print('Execution finished', u'\u2713')
```

Output:
```
The age of John Doe is 34
The age of Jane Doe is 44
Execution finished ✓
```

<!-- grade_id: tests-2-3, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import person_age_class
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from person_age_class import PersonAge
except:
    raise ValueError("Your solution does not contain the right function!")

import dataclasses
from datetime import date
from nose.tools import assert_equal

assert 'PersonAge' in dir(), 'Class PersonAge does not exist!'
assert dataclasses.is_dataclass(PersonAge), "The class is not a dataclass!"

p = PersonAge('John Doe', 'Manager', date(1990, 11, 15))
assert hasattr(p, 'name'), 'Class PersonAge does not have property `name`.'
assert hasattr(p, 'job'), 'Class PersonAge does not have property `job`.'
assert hasattr(p, 'age') == False, 'Class PersonAge should not have property `age`.'
p.compute_age()
assert hasattr(p, 'age'), 'Class PersonAge does not have property `age`.'
assert p.age == 34, f"The computed age `{p.age}` is incorrect!"

from unittest.mock import patch

with patch('__main__.print') as mock_print:
    s = str(p)
mock_print.assert_not_called()

person1 = PersonAge('John Doe', 'Manager', date(1990, 11, 15))
person1.compute_age()
person2 = PersonAge('Jane Doe', 'Manager', date(1980, 11, 15))
person2.compute_age()
assert person1.age < person2.age, "The age computation seems incorrect!"

validate_python_code("person_age_class.py")
print('tests passed', u'\u2713')
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

---
## Level 3: Pseudocode/HTML/CSS/JS problems

In these problems we are training the design part of the course, which will be used starting from week 6 in the course.

If the error messages are not very clear, you can use the online tool: https://onu1.s2.chalmers.se/nu/#textarea

Recommended sources:
- HTML: https://www.w3schools.com/html/
- CSS: https://www.w3schools.com/css/
- JavaScript: https://www.w3schools.com/js/

For this assignment, we study:
- lists: https://www.w3schools.com/html/html_lists.asp
- tables: https://www.w3schools.com/html/html_tables.asp

Visual Studio Code sometimes does not render HTML/CSS/JS as expected. We recommend testing your solutions here: https://www.w3schools.com/html/tryit.asp?filename=tryhtml_default



### Problem 3.1

Create a Python function called `list_generator` that receives a list of values.
The function should create a string that represents the HTML that shows the list of values, each one as an item of an unordered list.

The following example shows how the output of the test cell should look like.

Code for the test:
```
list_generator([
    'First item',
    'Second item',
    'Another item',
    'Final item'
])
```

Must result in the following result (the border is used just to delimit the output):
<html>
    <div style="border: 1px solid red;">
        <ul>
            <li>First item</li>
            <li>Second item</li>
            <li>Another item</li>
            <li>Final item</li>
        </ul>
    </div>
</html>

Note that each tag should be in one line, i.e., there should be a line break between every two tags.

```python
%%writefile list_generator_solution.py
### BEGIN SOLUTION
def list_generator(items: list[str]) -> str:
    # Create the opening tag for the unordered list
    html = "<ul>\n"
    # Add each item as a list item
    for item in items:
        html += f"  <li>{item}</li>\n"
    # Add the closing tag for the unordered list
    html += "</ul>"
    # Return the generated HTML string
    return html


### END SOLUTION

```

Output:
```
Overwriting list_generator_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from list_generator_solution import list_generator

print('HTML code:')
print('You can copy the code below and paste in the W3Schools tryit editor to test it!')
print(list_generator([
    'First item',
    'Second item',
    'Another item',
    'Final item'
]))
print('Execution finished', u'\u2713')
```

Output:
```
HTML code:
You can copy the code below and paste in the W3Schools tryit editor to test it!
<ul>
  <li>First item</li>
  <li>Second item</li>
  <li>Another item</li>
  <li>Final item</li>
</ul>
Execution finished ✓
```

```python
%autoreload 2
# DO NOT REMOVE THIS CODE >>> BEGINNING
# this cell assembles the HTML returned by the function that you wrote
# and renders it in the "preview" window
from IPython.display import display, HTML
from list_generator_solution import list_generator
html_header = """<!DOCTYPE html>
    <html lang="en">
        <head>
            <title>Page Title</title>
        </head>
        <body><p style="color: red;">start of your solution &gt;&gt;&gt;&gt;</p><hr/>
        """
html_footer = """     <hr/><p style="color: red;">&lt;&lt;&lt;&lt; end of your solution</p></body>
</html>"""

html_to_test = html_header + list_generator([
    'First item',
    'Second item',
    'Another item',
    'Final item'
]) + html_footer
display(HTML(html_to_test))

# DO NOT REMOVE THIS CODE <<< END
```

Output:
```
<IPython.core.display.HTML object>
```

```python
%autoreload 2
# test cell
from list_generator_solution import list_generator
from lorem_text import lorem
import random
from utils import validate_html

html_header = """<!DOCTYPE html>
    <html lang="en">
        <head>
            <title>Page Title</title>
        </head>
        <body>
        """
html_footer = """</body>
</html>"""

for i in range(5):

    num_items = random.randint(5, 20)

    items = []
    for i in range(num_items):
        items.append(lorem.words(random.randint(3, 20)))

    original_html = list_generator(items)

    html_to_test = html_header + original_html + html_footer

    soup = validate_html(html_to_test)

    # print(soup)

    tag_set = set([str(tag.name) for tag in soup.find_all()])
    tag_pairs = []
    for tag_1 in tag_set:
        for tag_2 in tag_set:
            if tag_1 != tag_2:
                # print(tag_1, tag_2)
                tag_pairs.append((tag_1, tag_2))

    for i, line in enumerate(original_html.split('\n')):
        for tag_1, tag_2 in tag_pairs:
            if f'<{tag_1}>' in line or f'</{tag_1}>' in line:
                assert f'<{tag_2}>' not in line and f'</{tag_2}>' not in line, f'Line {i+1} has two tags in the same line ({tag_1} and {tag_2}).\nYou must have them in different lines.\nLine {i+1}:\t{line}'

    tags_ul = soup.find_all('ul')
    assert len(tags_ul) == 1, f'The solution must have 1 unordered lists in your HTML. You have {len(tags_ul)} instead.'

    tags_li = tags_ul[0].find_all('li')
    assert len(tags_li) == num_items, f'The solution must have {num_items} list items in your HTML. You have {len(tags_li)} instead.'

    for idi, item in enumerate(tags_li):
        assert item.text == items[idi], f'The text in line {idi} is not the same as intended.'

validate_python_code("list_generator_solution.py")
print('tests passed', u'\u2713')
```

Output:
```
first
first
first
first
first
<IPython.core.display.HTML object>
tests passed ✓
```