# EEN060 - Applied object-oriented programming

Teacher: [Carlos Natalino](https://www.chalmers.se/en/staff/Pages/Carlos-Natalino-Da-Silva.aspx) / Examiner: [Paolo Monti](https://www.chalmers.se/en/staff/Pages/Paolo-Monti.aspx)

[Canvas course page](https://chalmers.instructure.com/courses/28094)

[Course channel on Chalmers Play](https://play.chalmers.se/channel/EEN060_EEN065_Applied_object-oriented_programming/300149)

Before you turn this assignment list in, make sure everything runs as expected.
First, **restart the kernel** and then **run all cells**.
Then, check if all the tests run correctly.
Note that if one of the problems present an error, the following ones **will not** be tested.

In case of discrepancies between the problem command and the tests, you should solve it having in mind the tests.

There are two types of cell:
1. *solution cells:* These are the cells where you write your answer, or modify the existing code to solve the problem.
2. *test cells:* These cells are used to test whether your solution is correct or not. If the tests run correctly, you should see a message `tests passed`. Otherwise, you should see an error message.

**Delete** the line `raise NotImplementedError()` from the problems that you solve.

**Do not delete or add any cell in this file.** All cells that you need are already in place.

You may also have problems that will be manually graded (e.g., HTML/CSS/JS/pseudocode problems).
These problems **do not** have a respective *test cell*.

If you want to execute a cell, select the cell and press **CTRL+Enter** (in Windows) or **CMD+Enter** (in macOS) or click on the **Run cell** button.

**If you want, you can solve this programming assignment using Google Colab**

Link: https://colab.research.google.com/

Just uncomment the following lines and run it.

```python
# !pip install --upgrade --user pip pip-tools
# !pip install ipykernel jupyter notebook wtforms flask flask-wtf gunicorn nose flask-testing pytest coverage flake8 requests mypy pylint blinker types-requests isort black lorem-text bs4 pep8-naming matplotlib redis
```

#### Assignment week 5

In this assignment, students should practice:
- how to intercept errors
- how to use exceptions to ensure input correctness
- how to handle multiple exception types
- how to open, read/write from/to files
- how to create a database
- continue getting familiar with Python and HTML

**Preparation:** Run the cell below every time you start working on this file, and every time you restart the kernel.

```python
%load_ext autoreload
try:
    from utils import validate_python_code, validate_html
except:
    import sys
    print("It seems this file is in the wrong folder. "
          "Make sure to place it in the `programming-assignments` folder/project.",
          file=sys.stderr)

def run_tests(func, test_cases):
    for (test_case, expected_output) in test_cases:
        assert func(test_case) == expected_output, \
            print("Your solution fails for the test case =", test_case, \
                "\nYour output is", func(test_case), \
                    "\nThe expected output is", expected_output)
```

## Level 1

Level 1 problems are simple problems that required only the knowledge obtained in the content of the lecture and the lecture notes of the particular lecture.

### Problem 1.1

Create a Python function called `validating_sum_of_squares` that receives a sequence of elements and validates whether or not each of the elements is a number and a sum of squares (eg. 25 is sum of 16 and 9 that is $4^2 + 3^2$).

Note that numbers can be placed within a string. And, for any integer $a$, since $a^2 = 0^2 + a^2$, square of any integer is assumed to valid.

At the end, the function should return a sequence (with the same number of elements as the input) containing `True` for the elements that are numbers which are sum of square and `False` otherwise.

**Examples**

Input: ['5', 3, 25, None]

Output: [True, False, True, False]

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-1, points: 0, grade: false, solution: true -->
```python
%%writefile validating_sum_of_squares_solution.py
# solution cell
### BEGIN SOLUTION
def validating_sum_of_squares(
    sequence: list[int | float | str | bool | None],
) -> list[bool]:
    def is_sum_of_squares(num: int | float) -> bool:
        # Validate whether a number can be expressed as a sum of two squares.
        if num < 0:  # Negative numbers cannot be sums of squares
            return False
        for i in range(int(num**0.5) + 1):  # Iterate over possible square roots
            j_squared = num - i**2
            if j_squared >= 0 and (j_squared**0.5).is_integer():
                return True
        return False

    results = []
    for element in sequence:
        # Convert to number if possible
        if isinstance(element, str) and element.isdigit():
            element = int(element)
        elif not isinstance(element, (int, float)):
            results.append(False)
            continue

        # Check if the current element is a sum of two squares
        if is_sum_of_squares(element):
            results.append(True)
        else:
            results.append(False)

    return results


### END SOLUTION

```

Output:
```
Overwriting validating_sum_of_squares_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from validating_sum_of_squares_solution import validating_sum_of_squares

seq_input = ['5', 3, 25, None, 'hello']
x = validating_sum_of_squares(seq_input)
print(f' Input sequence: ', '\t'.join([str(x) for x in seq_input]))
print(f"Output sequence: ", "\t".join([str(r) for r in x]))
print("Execution finished", u"\u2713")
```

Output:
```
Input sequence:  5	3	25	None	hello
Output sequence:  True	False	True	False	False
Execution finished ✓
```

<!-- grade_id: cell-tests-1-1, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import validating_sum_of_squares_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from validating_sum_of_squares_solution import validating_sum_of_squares
except:
    raise ValueError("Your solution does not contain the right function!")

test_cases = [(['2', 'x', -10, 3.3, 'asd', None, 'b', 4.0], \
    [True, False, False, False, False, False, False, True]), \
        ([9, 'x', -9, None], [True, False, False, False])]

run_tests(validating_sum_of_squares, test_cases)

validate_python_code("validating_sum_of_squares_solution.py")
print('tests passed', u'\u2713')
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 1.2

Create a Python function `gm_calculation` that receives a sequence of elements.

Then, the function should return the geometric mean of the numbers ignoring the elements in the sequence which are not numbers. The mean should be returned in a string type containing only two decimal cases. If there are no numbers in the sequence, then return `None`.

Remember to use the module `statistics` for this task.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-2, points: 0, grade: false, solution: true -->
```python
%%writefile gm_calculation_solution.py
# solution cell
### BEGIN SOLUTION
from statistics import geometric_mean


def gm_calculation(sequence: list[int | float | str | bool | None]) -> str | None:
    # Filter out elements that are not numbers
    numbers = []
    for x in sequence:
        if isinstance(x, (int, float)):
            numbers.append(x)
        elif isinstance(x, str):
            if x.isdigit():
                numbers.append(float(x))
            elif x.replace(".", "", 1).isdigit():
                numbers.append(float(x))
        else:
            continue
    print(numbers)
    # Check if there are no numbers in the list
    if not numbers:
        return None

    # Calculate the geometric mean using the `statistics` module
    gm = geometric_mean(numbers)

    # Return the result as a string with two decimal places
    return f"{gm:.2f}"


### END SOLUTION

```

Output:
```
Overwriting gm_calculation_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from gm_calculation_solution import gm_calculation

sequence = [10.2, 11, '12.3', 15, 32.5, 'x', '100.87534262']
print(f'The geometric mean of {sequence} is {gm_calculation(sequence)}')
print('Execution finished', u'\u2713')
```

Output:
```
[10.2, 11, 12.3, 15, 32.5, 100.87534262]
The geometric mean of [10.2, 11, '12.3', 15, 32.5, 'x', '100.87534262'] is 20.20
Execution finished ✓
```

<!-- grade_id: cell-test-1-2, points: 1, grade: true, solution: false -->
```python
%autoreload 2
#test cell
try:
    import gm_calculation_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from gm_calculation_solution import gm_calculation
except:
    raise ValueError("Your solution does not contain the right function!")

assert isinstance(gm_calculation([1.2, 1]), str), 'The return value should be a string'

test_cases = [([10.2, 11, 12.3, 15, 32.5, 'x', '100.87534262'], '20.20'), \
    (['10.2', 11, 12.3, 15, 32.5, None, '100.87534262'], '20.20'), \
    ([None, 'x', 'as', 'as1', 'x'], None), ([], None)]

run_tests(gm_calculation, test_cases)

validate_python_code("gm_calculation_solution.py")
print('tests passed', u'\u2713')
```

Output:
```
[1.2, 1]
[10.2, 11, 12.3, 15, 32.5, 100.87534262]
[10.2, 11, 12.3, 15, 32.5, 100.87534262]
[]
[]
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 1.3

Create a Python function called `datetime_validator` that receives a string which represents date-time, and converts it into `datetime` object.

The allowed format for the input: `Month Date, Year, Time`. For example `January 15, 2022, 12:30:45`. 

If the input is not in this format, return `None`.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-3, points: 0, grade: false, solution: true -->
```python
%%writefile datetime_validator_solution.py
# solution cell
### BEGIN SOLUTION
from datetime import datetime


def datetime_validator(date_time_str: str) -> datetime | None:
    try:
        # Define the required format
        format = "%B %d, %Y, %H:%M:%S"
        # Attempt to parse the string using the datetime.strptime method
        return datetime.strptime(date_time_str, format)
    except ValueError:
        # If parsing fails, return None
        return None


### END SOLUTION

```

Output:
```
Overwriting datetime_validator_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from datetime_validator_solution import datetime_validator

date_input = 'January 15, 2022, 12:30:45'
print(f'{date_input} is equivalent to {datetime_validator(date_input)} datetime object')
print('Execution finished', u'\u2713')
```

Output:
```
January 15, 2022, 12:30:45 is equivalent to 2022-01-15 12:30:45 datetime object
Execution finished ✓
```

<!-- grade_id: cell-test-1-3, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import datetime_validator_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from datetime_validator_solution import datetime_validator
except:
    raise ValueError("Your solution does not contain the right function!")

from datetime import datetime

test_cases = [('January 15, 2022, 12:30:45', datetime(2022, 1, 15, 12, 30, 45)),\
    ('September 08, 2021, 14:45:15', datetime(2021, 9, 8, 14, 45, 15)), \
    ('Sep 8, 2021, 14:45:15', None), ('2021-01-01 1:41:43.321', None)]

run_tests(datetime_validator, test_cases)

validate_python_code("datetime_validator_solution.py")
print('tests passed', u'\u2713')
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 1.4

Create a Python function called `write_text_with_line_numbers` that receives a string containing text and writes it to a text file called `text_file.txt`, with line numbers for each line. The function should always wipe out the previous content of the file, i.e., do not use the append mode.

**Example**

Input: "År\n2013"

Textfile contents: 
1. År
2. 2013

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-4, points: 0, grade: false, solution: true -->
```python
%%writefile write_text_with_line_numbers_solution.py
# solution cell
### BEGIN SOLUTION
def write_text_with_line_numbers(text: str) -> None:
    # Split the input text into lines
    lines = text.splitlines()

    # Open the file `text_file.txt` in write mode ('w') to overwrite existing content
    with open("text_file.txt", "w", encoding="utf-8") as file:
        # Loop through each line along with its line number
        for i, line in enumerate(lines, start=1):
            # Write the line number followed by the text line to the file
            if i < len(lines):
                file.write(f"{i}. {line}\n")
            else:
                file.write(f"{i}. {line}")


### END SOLUTION

```

Output:
```
Overwriting write_text_with_line_numbers_solution.py
```

```python
%autoreload 2
# uncomment the line below to debug
from write_text_with_line_numbers_solution import write_text_with_line_numbers

write_text_with_line_numbers("År\n2013")
print('The file contents are:', u'\u2713')
with open("text_file.txt", "r", encoding="utf-8") as file:
    print(file.read())
```

Output:
```
The file contents are: ✓
1. År
2. 2013
```

<!-- grade_id: cell-test-1-4, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import write_text_with_line_numbers_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from write_text_with_line_numbers_solution import write_text_with_line_numbers
except:
    raise ValueError("Your solution does not contain the right function!")

# test cell
from nose.tools import assert_equal
import os

if os.path.isfile('text_file.txt'):
    os.remove('text_file.txt')

write_text_with_line_numbers("År\n2013")
assert os.path.isfile('text_file.txt'), 'The file was not created!'
with open("text_file.txt", "r", encoding="utf-8") as file:
    assert file.read() == "1. År\n2. 2013", "The contents of the file are incorrect!"
os.remove('text_file.txt')

write_text_with_line_numbers("one\ntwo\nthree\nfour")
assert os.path.isfile('text_file.txt'), 'The file was not created!'
with open("text_file.txt", "r", encoding="utf-8") as file:
    assert file.read() == "1. one\n2. two\n3. three\n4. four", \
        "The contents of the file are incorrect!"
os.remove('text_file.txt')

validate_python_code("write_text_with_line_numbers_solution.py")
print('tests passed', u'\u2713')
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

---
## Level 2

Problems in level 2 should build upon the content of the lecture and the content of previous lectures.
Therefore, the problems can be solved based solely on the lecture notes available for the course up to the lecture the assignment is about.

### Problem 2.1

Create a Python class named `Student` that is also a dataclass.
The attributes are:
- name, string
- birthdate, date
- person_number, string
- phone_number, string (optional, default `None`)
- address, string (optional, default `None`)

The solution should consider receiving a date object for birthdate.


```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: answer-1-6, points: 0, grade: false, solution: true -->
```python
%%writefile student_class.py
# solution cell
### BEGIN SOLUTION
from dataclasses import dataclass, field
from datetime import date
from typing import Optional


@dataclass
class Student:
    name: str
    birthdate: date
    person_number: str
    phone_number: Optional[str] = field(default=None)
    address: Optional[str] = field(default=None)


### END SOLUTION

```

Output:
```
Overwriting student_class.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from datetime import date
from student_class import Student

student = Student('test', date(1998, 3, 4), '1124234', '12314234', 'Gibraltargatan')
print(f'The Student object converts to string as \n{student}')
print('Execution finished', u'\u2713')
```

Output:
```
The Student object converts to string as 
Student(name='test', birthdate=datetime.date(1998, 3, 4), person_number='1124234', phone_number='12314234', address='Gibraltargatan')
Execution finished ✓
```

<!-- grade_id: tests-1-6, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import student_class
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from student_class import Student
except:
    raise ValueError("Your solution does not contain the right function!")

from dataclasses import is_dataclass

assert 'Student' in dir(), 'Class Student does not exist!'
assert is_dataclass(Student), "The Student class is not a dataclass!"

validate_python_code("student_class.py")
print('tests passed', u'\u2713')
```

Output:
```
<IPython.core.display.HTML object>
tests passed ✓
```

### Problem 2.2

Create a Python function called `insert_student` that receives the following parameters:

- name, string
- birthdate, date
- person_number, string
- phone_number, string (optional, default `None`)
- address, string (optional, default `None`)

Then, the function creates an object of the class `Student` and persists it to the database, returning how many objects of the class `Student` exist in the database.
The function should use the *key* `dataset` to retrieve and store the list of `Student` objects in the database.

The solution should consider receiving a date object for birthdate.

If a `Student` object with the same person_number already exists, the function should not insert the new `Student` object.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
%%writefile database_setup.py
# EXECUTE TO CREATE THE NECESSARY FILE
from redis import Redis

# configuration
DATABASE_HOST="onu1.s2.chalmers.se"
DATABASE_PORT=6380
DATABASE_DB=1
DATABASE_PASSWORD="6a580447-cc21-422e-8d74-66f93bdee3ab"

# connecting
db: Redis = Redis(
    host=DATABASE_HOST,
    port=DATABASE_PORT,
    db=DATABASE_DB,
    password=DATABASE_PASSWORD
)
db.ping()
db.delete("dataset")
```

Output:
```
Overwriting database_setup.py
```

<!-- grade_id: cell-answer-2-2, points: 0, grade: false, solution: true -->
```python
%%writefile insert_student_solution.py
# solution cell
### BEGIN SOLUTION
from datetime import date


def insert_student(
    name: str,
    birthdate: date,
    person_number: str,
    phone_number: Optional[str] = None,
    address: Optional[str] = None,
) -> int:
    
    # Retrieve the current dataset from the global "database"
    dataset: List[Student] = database["dataset"]
    
    # Check for an existing Student with the same person_number
    for student in dataset:
        if student.person_number == person_number:
            # If a match is found, no insertion occurs, return current count
            return len(dataset)
    
    # Create and insert new Student
    new_student = Student(
        name=name,
        birthdate=birthdate,
        person_number=person_number,
        phone_number=phone_number,
        address=address,
    )
    dataset.append(new_student)

    # Update the database
    database["dataset"] = dataset

    # Return the total number of Students
    return len(dataset)


### END SOLUTION

```

Output:
```
Overwriting insert_student_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from datetime import date

from database_setup import db
from insert_student_solution import insert_student
from student_class import Student

db.delete("dataset")

cid = insert_student('Jane Doe', date(1998, 3, 4), '9803041234', '781452123', 'ABC')
print(f'We currently have {cid} `Student` objects in the database!', u'\u2713')

cid = insert_student('Jane Doe', date(1998, 3, 4), '9803041234', '781452123', 'ABC')
print(f'We currently have {cid} `Student` objects in the database!', u'\u2713')

cid = insert_student('Jack K', date(1993, 1, 7), '9301072345', '72345134', 'XYZ')
print(f'We currently have {cid} `Student` objects in the database!', u'\u2713')
```

Output:
```
File ~/Library/CloudStorage/OneDrive-Chalmers/Documents/STUDY-PHD/04_Teaching_Assistant/2025/EEN060_Applied_object-oriented_programming/programming-assignments/.venv/lib/python3.10/site-packages/redis/_parsers/resp2.py:15, in _RESP2Parser.read_response(self, disable_decoding)
     13 pos = self._buffer.get_pos() if self._buffer else None
     14 try:
---> 15     result = self._read_response(disable_decoding=disable_decoding)
     16 except BaseException:
     17     if self._buffer:

File ~/Library/CloudStorage/OneDrive-Chalmers/Documents/STUDY-PHD/04_Teaching_Assistant/2025/EEN060_Applied_object-oriented_programming/programming-assignments/.venv/lib/python3.10/site-packages/redis/_parsers/resp2.py:38, in _RESP2Parser._read_response(self, disable_decoding)
     35 # if the error is a ConnectionError, raise immediately so the user
     36 # is notified
     37 if isinstance(error, ConnectionError):
---> 38     raise error
     39 # otherwise, we're dealing with a ResponseError that migh
```

<!-- grade_id: cell-test-2-2, points: 1, grade: true, solution: false -->
```python
%autoreload 2
# test cell
try:
    import insert_student_solution
except:
    raise ValueError("You did not execute your solution cell!")
try:
    from insert_student_solution import insert_student
except:
    raise ValueError("Your solution does not contain the right function!") 

from datetime import date
from database_setup import db

db.delete("dataset")
assert insert_student('Jane Doe', date(1998, 3, 4), '9803041234', '781452123', 'ABC') \
    == 1, 'Student object was not persisted successfully!'
assert insert_student('Jane Doe', date(1998, 3, 4), '9803041234', '781452123', 'ABC') \
    == 1, 'Expecting one object, but got a different number!'
assert insert_student('Jack K', date(1993, 1, 7), '9301072345', '72345134', 'XYZ') \
    == 2, 'Student object was not persisted successfully!'

validate_python_code("insert_student_solution.py")
print('tests passed', u'\u2713')
```

Output:
```
ValueError                                Traceback (most recent call last)
Cell In[56], line 6
      4     import insert_student_solution
      5 except:
----> 6     raise ValueError("You did not execute your solution cell!")
      7 try:
      8     from insert_student_solution import insert_student

ValueError: You did not execute your solution cell!
```

---
## Level 3: Pseudocode/HTML/CSS/JS problems

In these problems we are training the design part of the course, which will be used starting from week 6 in the course.

If the error messages are not very clear, you can use the online tool: https://onu2.s2.chalmers.se/nu/#textarea

Recommended sources:
- HTML: https://www.w3schools.com/html/
- CSS: https://www.w3schools.com/css/
- JavaScript: https://www.w3schools.com/js/

For this assignment, we study:
- tables: https://www.w3schools.com/html/html_tables.asp

Visual Studio Code sometimes does not render HTML/CSS/JS as expected. We recommend testing your solutions here: https://www.w3schools.com/html/tryit.asp?filename=tryhtml_default



### Problem 3.1

Create a Python function called `table_generator` that receives a list.

The function should return a string that represents the HTML code. The HTML code represents the input list as a table with two columns: 
- index in the list, and
- value at that index.

Index and value should be explicitly defined as the heads of the table.

The following example shows how the output of the test cell should look like.

Code for testing:
```
table_generator(
    ['test value 1', 
    'test value 2', 
    'test value 3', 
    'test value 4', 
    'test value 5', 
    'test value 6']
    )
```

Must result in the following output  (the border is used just to delimit the output):

<div style="border: 1px solid red;">
<table>
    <thead>
        <tr>
            <th>Index</th> <th>Value</th>
        </tr>
    </thead>
    <tbody>
        <tr><td>1</td><td>test value 1</td></tr>
        <tr><td>2</td><td>test value 2</td></tr>
        <tr><td>3</td><td>test value 3</td></tr>
        <tr><td>7</td><td>test value 4</td></tr>
        <tr><td>5</td><td>test value 5</td></tr>
        <tr><td>6</td><td>test value 6</td></tr>
    </tbody>
</table>
</div>

Note that each tag should be in one line, i.e., there should be a line break between every two tags.

```python
%%writefile table_generator_solution.py
# solution cell
### BEGIN SOLUTION
def table_generator(input_list: list[str]) -> str:
    # Start the HTML table
    html = "<table>\n"
    # Add table head with <thead>
    html += "  <thead>\n"
    html += "    <tr>\n"
    html += "      <th>Index</th>\n"
    html += "      <th>Value</th>\n"
    html += "    </tr>\n"
    html += "  </thead>\n"
    # Add table body with <tbody>
    html += "  <tbody>\n"
    for index, value in enumerate(input_list, start=1):
        html += "    <tr>\n"
        html += f"      <td>{index}</td>\n"
        html += f"      <td>{value}</td>\n"
        html += "    </tr>\n"
    html += "  </tbody>\n"
    # Close the table
    html += "</table>"
    return html


### END SOLUTION

```

Output:
```
Overwriting table_generator_solution.py
```

```python
%autoreload 2
# uncomment the line(s) below to debug
from table_generator_solution import table_generator

print('HTML code:')
s = table_generator(
    ['test value 1', 
    'test value 2', 
    'test value 3', 
    'test value 4', 
    'test value 5', 
    'test value 6']
    )
print(s)
print('Execution finished', u'\u2713')
```

Output:
```
HTML code:
<table>
  <thead>
    <tr>
      <th>Index</th>
      <th>Value</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>1</td>
      <td>test value 1</td>
    </tr>
    <tr>
      <td>2</td>
      <td>test value 2</td>
    </tr>
    <tr>
      <td>3</td>
      <td>test value 3</td>
    </tr>
    <tr>
      <td>4</td>
      <td>test value 4</td>
    </tr>
    <tr>
      <td>5</td>
      <td>test value 5</td>
    </tr>
    <tr>
      <td>6</td>
      <td>test value 6</td>
    </tr>
  </tbody>
</table>
Execution finished ✓
```

```python
%autoreload 2
# DO NOT REMOVE THIS CODE >>> BEGINNING
# this cell assembles the HTML returned by the function that you wrote
# and renders it in the "preview" window

#from IPython import display
from table_generator_solution import table_generator
from IPython.display import display, HTML
html_header = """<!DOCTYPE html>
    <html lang="en">
        <head>
            <title>Page Title</title>
        </head>
        <body><p style="color: red;">start of your solution &gt;&gt;&gt;&gt;</p>
        """
html_footer = """<p style="color: red;">&lt;&lt;&lt;&lt; end of your solution</p></body>
</html>"""
html = table_generator(
    ['test value 1', 
    'test value 2', 
    'test value 3', 
    'test value 4', 
    'test value 5', 
    'test value 6']
    )
html_to_test = html_header + html + html_footer
display(HTML(html_to_test))

# DO NOT REMOVE THIS CODE <<< END
```

Output:
```
<IPython.core.display.HTML object>
```

```python
%autoreload 2
# test cell
import random
from lorem_text import lorem
from bs4 import BeautifulSoup
from table_generator_solution import table_generator

has_error = False

html_header = """<!DOCTYPE html>
    <html lang="en">
        <head>
            <title>Page Title</title>
        </head>
        <body><p style="color: red;">start of your solution &gt;&gt;&gt;&gt;</p>
        """
html_footer = """<p style="color: red;">&lt;&lt;&lt;&lt; end of your solution</p></body>
</html>"""

for i in range(5):

    input_list = []
    num_elems = random.randint(5, 10)
    for i in range(num_elems):
        input_list.append(lorem.words(random.randint(5, 10)))

    original_html = table_generator(input_list)

    html_to_test = html_header + original_html + html_footer

    soup = validate_html(html_to_test)

    tag_set = set([str(tag.name) for tag in soup.find_all()])
    tag_pairs = []
    for tag_1 in tag_set:
        for tag_2 in tag_set:
            if tag_1 != tag_2:
                tag_pairs.append((tag_1, tag_2))

    for i, line in enumerate(original_html.split('\n')):
        for tag_1, tag_2 in tag_pairs:
            if f'<{tag_1}>' in line or f'</{tag_1}>' in line:
                assert f'<{tag_2}>' not in line and f'</{tag_2}>' not in line, \
                    f'Line {i+1} has two tags in the same line ({tag_1} and {tag_2}).\
                        \nYou must have them in different lines.\nLine {i+1}:\t{line}'

    tags_table = soup.find_all('table')
    assert len(tags_table) == 1, f'The HTML seems needs to have 1 table. Now it has \
        {len(tags_table)} table(s).'

    tags_thead = tags_table[0].find_all('thead')
    assert len(tags_thead) == 1, 'The needs to have one `thead` tag.'

    tags_tr = tags_thead[0].find_all('tr')
    assert len(tags_tr) == 1, 'The table head needs to have only one row.'

    tags_th = tags_tr[0].find_all('th')
    assert len(tags_th) == 2, f'The table must have 2 header columns. Now it has \
        {len(tags_th)} column(s).'

    tags_tbody = tags_table[0].find_all('tbody')
    assert len(tags_tbody) == 1, 'Your table is missing a `tbody` tag.'

    tags_tr = tags_tbody[0].find_all('tr')
    assert len(tags_tr) == num_elems, f'The table must have {num_elems} rows. \
        Now it has {len(tags_tr)} row(s).'

    for idx, tag_tr in enumerate(tags_tr):
        tags_td = tag_tr.find_all('td')
        assert len(tags_td) == 2, f'The table must have 2 columns. Now it has \
            {len(tags_td)} column(s).\nRow with problems: \t{tag_tr}'
        assert int(tags_td[0].contents[0]) == idx + 1, f'The table contents are \
            incorrect. Row with problems: \t{tag_tr}'
        assert tags_td[1].contents[0] == input_list[idx], f'The table contents are \
            incorrect. Row with problems: \t{tag_tr}'

validate_python_code("table_generator_solution.py")
print('tests passed', u'\u2713')
```

Output:
```
first
first
first
first
first
<IPython.core.display.HTML object>
tests passed ✓
```